# Imports

The first thing we'll need to do is load in the libraries we'll be using. 

```python
import numpy as np
import pandas as pd
import seaborn as sb
import matplotlib.pyplot as plt

import sklearn
from sklearn.preprocessing import MinMaxScaler
from sklearn.preprocessing import MaxAbsScaler
from sklearn.preprocessing import StandardScaler
from sklearn.preprocessing import RobustScaler
from sklearn.preprocessing import Normalizer
from sklearn.preprocessing import QuantileTransformer
from sklearn.preprocessing import PowerTransformer

#sb.set_theme(style="darkgrid")
#np.random.seed(0) # set seed for reproducibility

print("Numpy:", np.__version__)
print("Pandas:", pd.__version__)
print("Seaborn:", sb.__version__)
```

# Generate data

```python
N_SAMPLES = 10000
BINS      = 30
npr       = np.random.RandomState(304)
```

<!-- grade_id: cell-73b2184af646ab0a, points: 0, grade: false, solution: true -->
```python
#2 points 
# create a dataframe that contains a columns for all the next normalization and scaling methods
# Exponential
# Lognormal
# Chi-sqaured
# Weibull
# Gaussian
# Uniform
# Bimodal

############### DELETE THE RAISE NOTIMPLEMENTED ERROR LINES!!!!!!

df = None

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-ae7968d6618817d8, points: 2, grade: true, solution: false -->
```python
assert df["Exponential"].shape[0] == 10000
assert df["Lognormal"].shape[0] == 10000
assert df["Chi-squared"].shape[0] == 10000
assert df["Weibull"].shape[0] == 10000
assert df["Gaussian"].shape[0] == 10000
assert df["Uniform"].shape[0] == 10000
assert df["Bimodal"].shape[0] == 10000
assert df.shape == (10000, 7)
```

```python
df.hist(bins=BINS, layout=(3, 4), figsize=(16,8));
```

# Scaling vs. Normalization: What's the difference?

One of the reasons that it's easy to get confused between scaling and normalization is because the terms are sometimes used interchangeably and, to make it even more confusing, they are very similar! In both cases, you're transforming the values of numeric variables so that the transformed data points have specific helpful properties. The difference is that:
- in **scaling**, you're changing the *range* of your data, while 
- in **normalization**, you're changing the *shape of the distribution* of your data. 

Let's talk a little more in-depth about each of these options. 

# Scaling

This means that you're transforming your data so that it fits within a specific scale, like 0-100 or 0-1.  You want to scale data when you're using methods based on measures of how far apart data points are, like [support vector machines (SVM)](https://en.wikipedia.org/wiki/Support_vector_machine) or [k-nearest neighbors (KNN)](https://en.wikipedia.org/wiki/K-nearest_neighbors_algorithm). With these algorithms, a change of "1" in any numeric feature is given the same importance. 

For example, you might be looking at the prices of some products in both Yen and US Dollars. One US Dollar is worth about 100 Yen, but if you don't scale your prices, methods like SVM or KNN will consider a difference in price of 1 Yen as important as a difference of 1 US Dollar! This clearly doesn't fit with our intuitions of the world. With currency, you can convert between currencies. But what about if you're looking at something like height and weight? It's not entirely clear how many pounds should equal one inch (or how many kilograms should equal one meter).

By scaling your variables, you can help compare different variables on equal footing. To help solidify what scaling looks like, let's look at a made-up example.

<!-- grade_id: cell-7b06b3bfbc458f9c, points: 0, grade: false, solution: true -->
```python
# 1 point
# use the min max scaler to scale your your data
original_data = df[["Exponential"]]

scaled_data = None
# YOUR CODE HERE
raise NotImplementedError()
# plot both together to compare
fig, ax = plt.subplots(1,2, figsize=(12,4))
sb.histplot(original_data, ax=ax[0]), ax[0].set_title("Original Data")
sb.histplot(scaled_data,   ax=ax[1]), ax[1].set_title("Scaled data");
```

<!-- grade_id: cell-cf85d2040d8c1d8e, points: 1, grade: true, solution: false -->
```python
assert scaled_data[scaled_data.argmax()] <= 1
assert scaled_data[scaled_data.argmin()] >= 0
```

Notice that the *shape* of the data doesn't change, but that instead of ranging from 0 to 8ish, it now ranges from 0 to 1.

# Normalization

Scaling just changes the range of your data. Normalization is a more radical transformation. The point of normalization is to change your observations so that they can be described as a normal distribution.

> **[Normal distribution:](https://en.wikipedia.org/wiki/Normal_distribution)** Also known as the "bell curve", this is a specific statistical distribution where a roughly equal observations fall above and below the mean, the mean and the median are the same, and there are more observations closer to the mean. The normal distribution is also known as the Gaussian distribution.

In general, you'll normalize your data if you're going to be using a machine learning or statistics technique that assumes your data is normally distributed. Some examples of these include linear discriminant analysis (LDA) and Gaussian naive Bayes. (Pro tip: any method with "Gaussian" in the name probably assumes normality.)

The method we're using to normalize here is called the [Box-Cox Transformation](https://en.wikipedia.org/wiki/Power_transform#Box%E2%80%93Cox_transformation). Let's take a quick peek at what normalizing some data looks like:

<!-- grade_id: cell-94671ef831c6a242, points: 0, grade: false, solution: true -->
```python
# 1 point

# normalize the exponential data with boxcox

normalized_data = None

# YOUR CODE HERE
raise NotImplementedError()
fig, ax = plt.subplots(1,2, figsize=(12,4))
sb.histplot(original_data,    ax=ax[0]), ax[0].set_title("Original Data")
sb.histplot(normalized_data,  ax=ax[1]), ax[1].set_title("Normalized data");
```

<!-- grade_id: cell-68159b6d3588d1c8, points: 1, grade: true, solution: false -->
```python
assert normalized_data[normalized_data.argmax()] < 10
assert normalized_data[normalized_data.argmin()] > -7
```

Notice that the *shape* of our data has changed. Before normalizing it was almost L-shaped. But after normalizing it looks more like the outline of a bell (hence "bell curve"). 

<!-- grade_id: cell-fd73d602ff0531f7, points: 0, grade: false, solution: true -->
```python
# 1 point
# Get the skewness of all the samples of the original dataframe
skew = None

# YOUR CODE HERE
raise NotImplementedError()
skew
```

<!-- grade_id: cell-3cbcbdf2330141f3, points: 1, grade: true, solution: false -->
```python
assert skew[0] > 1
assert skew[1] > 3
assert skew[2] > 1
assert -1 < skew[3] < 1
assert -1 < skew[4] < 1
assert -1 < skew[5] < 1
assert -1 < skew[6] < 1
```

![](https://www.allaboutcircuits.com/uploads/articles/understanding-the-normal-distribution-parametric-tests-skewness-and-kurtosis-rk-aac-image2.jpg)

<!-- grade_id: cell-f8b7d87f64909213, points: 0, grade: false, solution: true -->
```python
from scipy.special import erfinv
# 5 points
# implement the rank gaussian fucntion from the slides
def rank_gauss(x):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-362e7a1105baf21d, points: 5, grade: true, solution: false -->
```python
gauss = rank_gauss(df.Uniform)
assert gauss.shape[0] == 10000
assert gauss[gauss.argmax()] <= 4
assert gauss[gauss.argmin()] >= -4
```

```python
gauss.plot.hist()
```

<!-- grade_id: cell-15b6c59332351df4, points: 0, grade: false, solution: true -->
```python
# 3 points
# create a dictionary with the following enconders
'''
encoders = {
  # Scalers
  "Min-max scaling":   
  "Max-abs scaling":    
  "Standard scaling":   
  "Robust scaling":     
    
  # Normalizers
  "Box-Cox":                # Can only be used for positive values
  "Yeo-Johnson":         # Similar to Box-cox but can be used for negative values.
  "Quantile (normal)":  
  "Quantile (uniform)": 
}
'''
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-bd44d6b1c55ae2c1, points: 3, grade: true, solution: false -->
```python
assert type(encoders["Min-max scaling"]) is sklearn.preprocessing._data.MinMaxScaler
assert type(encoders["Max-abs scaling"]) is sklearn.preprocessing._data.MaxAbsScaler
assert type(encoders["Standard scaling"]) is sklearn.preprocessing._data.StandardScaler
assert type(encoders["Robust scaling"]) is sklearn.preprocessing._data.RobustScaler

assert type(encoders["Box-Cox"]) is sklearn.preprocessing._data.PowerTransformer
assert type(encoders["Yeo-Johnson"]) is sklearn.preprocessing._data.PowerTransformer
assert type(encoders["Quantile (normal)"]) is sklearn.preprocessing._data.QuantileTransformer
assert type(encoders["Quantile (uniform)"]) is sklearn.preprocessing._data.QuantileTransformer
```

```python
print("Original distributions")
df.hist(bins=BINS, layout=(1,7), figsize=(16,2))
plt.show()
for encoder_name, encoder in encoders.items():
    print(encoder_name)
    df_tmp = pd.DataFrame(encoder.fit_transform(df), columns=df.columns)
    df_tmp.hist(bins=BINS, layout=(1,7), figsize=(16,2))
    plt.show()
```

```python

```



# Your turn

It's time to [**apply what you just learned**](https://www.kaggle.com/kernels/fork/10824404) a dataset of Kickstarter projects.