```python
import pandas as pd
import matplotlib.pyplot as plt
#import seaborn as sb
import sklearn as skl
import time
from IPython.display import display, clear_output

import sklearn
from sklearn import pipeline      # Pipeline
from sklearn import preprocessing # OrdinalEncoder, LabelEncoder
from sklearn import impute
from sklearn import compose
from sklearn import model_selection # train_test_split
from sklearn import metrics         # accuracy_score, balanced_accuracy_score, plot_confusion_matrix
from sklearn import set_config

set_config(display='diagram') # Useful for display the pipeline

print("Pandas  ", pd.__version__)
print("Sklearn ", skl.__version__) # Try to use 0.24
```

```python
path = "./tabular-playground-feb21/"
!ls $path
```

### Load data

<!-- grade_id: cell-52a8794a601f4d0e, points: 0, grade: false, solution: true -->
```python
# 1 point 
 #load the train, test and sample_submission data using the is as the index column
df = None
df_t = None
sub = None

# YOUR CODE HERE
raise NotImplementedError()

print("Train dataframe:", df.shape)
print("Test dataframe:", df_t.shape)
df.head()
```

<!-- grade_id: cell-63c54563653a95db, points: 1, grade: true, solution: false -->
```python
print(sub.shape)
assert df.shape == (300000,25) 
assert df_t.shape == (200000,24) 
assert sub.shape == (200000,1) 
```

### Subsample the data

<!-- grade_id: cell-af0fe9c9d53bb080, points: 0, grade: false, solution: true -->
```python
#1 point
# get a random sample of size 50k, tested with random state 0
dfs = None
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-6d044ab7181322e1, points: 1, grade: true, solution: false -->
```python
assert dfs.shape == (50000,25)
```

### There are no missings in this dataset

```python
df.isna().sum().sum(), df_t.isna().sum().sum()
```

```python
df.columns
```

### Define variables for x and y

<!-- grade_id: cell-798649ceac290b47, points: 0, grade: false, solution: true -->
```python
cat_vars = ['cat0', 'cat1', 'cat2', 'cat3', 'cat4', 'cat5', 'cat6', 'cat7', 'cat8', 'cat9']
num_vars = ['cont0', 'cont1', 'cont2', 'cont3',  'cont4',  'cont5',  'cont6',
            'cont7', 'cont8', 'cont9', 'cont10', 'cont11', 'cont12', 'cont13']

# 1 point 
# store in X all the categorical and numerical columns
# store in Y the target

x = 0
y = 0
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-a12e77048aa686c4, points: 1, grade: true, solution: false -->
```python
assert x.shape == (50000,24)
assert y.shape == (50000,)
```

### See the target variable ([The metric is RMSE](https://www.kaggle.com/c/tabular-playground-series-feb-2021/overview/evaluation))

```python
df.target.plot.hist(bins=500);
```

### See the numerical variables

```python
for var in num_vars:
    df[var].plot.hist(bins=500)
    plt.show();
```

### See the categorical variables (ypu can use `bar()`, `barh()` or `.pie()`)

```python
for var in cat_vars:
    #df[var].value_counts().plot.bar()
    #df[var].value_counts().plot.barh()
    df[var].value_counts().plot.pie()
    plt.show();
```

# Exercise 1 (2 pts)

Create the `ColumTransformer` for **models based on multiplications and distances**:
- Pipeline for numerical variables (`num_vars`): `Standardscaler()` or `QuantileTransformer(n_quantiles=100, output_distribution='normal')`
- Pipeline for categorical variables (`cat_vars`): `OneHotEncoder(handle_unknown='ignore')`

<!-- grade_id: cell-1bcb9c03884767e0, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

tree_prepro
```

<!-- grade_id: cell-25faeb14cfb1e048, points: 2, grade: true, solution: false -->
```python
assert type(tree_prepro) is sklearn.compose._column_transformer.ColumnTransformer
```

# Exercise 2 (2 pts)

Create the `ColumTransformer` for **models based on trees**:
- Pipeline for numerical variables (`num_vars`): **Do nothing**. Tip: You can use `remainder='passthrough'` when you create the `ColumTransformer`.
- Pipeline for categorical variables (`cat_vars`): `OrdinalEncoder(handle_unknown='use_encoded_value', unknown_value=-1)[source]`

<!-- grade_id: cell-d5bc34b399dc8b7a, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

mult_prepro
```

<!-- grade_id: cell-5938715b631f6081, points: 2, grade: true, solution: false -->
```python
assert type(mult_prepro) is sklearn.compose._column_transformer.ColumnTransformer
```

## Exercise 3 (2pts)
Reasearch about som models and its hyperparamters

<!-- grade_id: cell-b3793d6c01aa98b1, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-b2aa9f758935ebbf, points: 0, grade: false, solution: true -->
```python
x_train, x_valid, y_train, y_valid = model_selection.train_test_split(x, y, test_size=0.4, random_state=0)

results = pd.DataFrame({'Model': [], 'MSE': [], 'RMSE': [], 'Time': []})

for name, pipe in all_pipelines.items():

    start_time = time.time()
    preds = 0
    # YOUR CODE HERE
    raise NotImplementedError()
    
    """
    results = results.append({"Model": name,
                              "MSE":   # mean sqaure error,
                              "RMSE":  #mean square error squared,
                              "Time":  total_time},
                              ignore_index=True)
    """
    
    results_ord = results.sort_values(by=['RMSE'], ascending=True, ignore_index=True)
    results_ord.index += 1 
    clear_output()
    display(results_ord.style.bar(subset=['MSE', 'RMSE'], vmin=0, color='#5fba7d'))
```

<!-- grade_id: cell-56489e8dd6257a2a, points: 3, grade: true, solution: false -->
```python
assert results_ord['MSE'].max() < 2
assert results_ord['MSE'].min() < 0.8
assert results_ord['RMSE'].max() < 1.5
assert results_ord['RMSE'].min() < 0.9
```

### Optional Exercise (optional because it takes a lot of time of computation)
Train with a more realable validation data the best models (maybe the best 5 models). You can:
- Use all training data (without subsample)
- Use Crosvalidation (`KFold`)(using 3, 5 or 7 folds) instead of `train_test_split`

```python
df   = pd.read_csv(path+"train.csv", index_col="id")

print("Using all data with", len(df), "rows")
x = df[num_vars+cat_vars]
y = df["target"]
```

<!-- grade_id: cell-2f06f87cae93f9e9, points: 0, grade: false, solution: true -->
```python
# IMPLEMENT SOMETHING SIMILAR TO THE PROCESS BEFORE
# BUT USE KFOLD

# YOUR CODE HERE
raise NotImplementedError()
```

```python

```