# for Preprocessing NLP Data
- removing non alphabatic characters
- lowering the format
- tokeninzation
- remove stop words
- steming / lemitization
- spell correction


```python
from nltk.tokenize import word_tokenize 
from nltk.corpus import stopwords
from nltk.stem import PorterStemmer
import nltk
stemmer=PorterStemmer()

from wordcloud import WordCloud

import pandas as pd
import matplotlib.pyplot as plt
import numpy as np

from autocorrect import Speller
spell=Speller()
from tqdm import tqdm
import re

```

<!-- grade_id: cell-331962f1b50a615b, points: 0, grade: false, solution: true -->
```python
# 1 point
# load the data from spam.csv
data = 0
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-c04c4903d286f150, points: 1, grade: true, solution: false -->
```python
assert data.shape == (5572,5)
```

```python
data.head()
```

<!-- grade_id: cell-f632d08142367603, points: 0, grade: false, solution: true -->
```python
# 1 point
# drop the columns unnamed 2,3,4

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-9a447414f829860c, points: 1, grade: true, solution: false -->
```python
assert data.shape == (5572,2)
```

```python
data.head()
```

<!-- grade_id: cell-cf617156831b7c5d, points: 0, grade: false, solution: true -->
```python
# 1 POINT
# rename the columns to labels and messages
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-38fc37d207bd59d8, points: 1, grade: true, solution: false -->
```python
assert data['labels'].shape == (5572,)
assert data['messages'].shape == (5572,)
```

```python
data.head()
```

# Lets Visualize the Spam words

```python
spam_words= ' '.join(list(data[data['labels']=='spam']['messages']))
spam_wc= WordCloud(width=500,height=500).generate(spam_words)
plt.figure(figsize=(10,8))
plt.imshow(spam_wc)
plt.show()
```

# Lets Visualize the Ham words

```python
spam_words= ' '.join(list(data[data['labels']=='ham']['messages']))
spam_wc= WordCloud(width=500,height=500).generate(spam_words)
plt.figure(figsize=(10,8))
plt.imshow(spam_wc)
plt.show()
```

# Preprocessing the Data

<!-- grade_id: cell-79b3250b45413841, points: 0, grade: false, solution: true -->
```python
# 2 points 
# change the labels to 1 or 0 instead of 'spam' and 'ham'
# Encoding our Labels

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-d6790d60a9b557e0, points: 2, grade: true, solution: false -->
```python
assert data['labels'][0] == 0
assert data['labels'][1] == 0
assert data['labels'][2] == 1
```

```python
# this part is given to you as a quick glance on the preprocessing in NLP

# preprocessing the Messages

nltk.download('punkt')
nltk.download('stopwords')

new_data=[]

for i in tqdm(range(data.shape[0])):
    lines=data.iloc[i,1]
    # removing non alphabatic characters
    lines= re.sub('[^A-Za-z]',' ',lines)
    # lowering the every word
    lines=lines.lower()
    
    # tokenization
    tokenized_lines=word_tokenize(lines)
    
    # removing stop words ,stemming and spell correction
    processed_lines=[]
    for i in tokenized_lines:
        if i not in set(stopwords.words('english')):
            processed_lines.append(spell(stemmer.stem(i)))
            
    final_lines=' '.join(processed_lines)
    new_data.append(final_lines)
    
            
    
```

```python
new_data
```

```python
len(new_data)
```

```python
Y=data['labels']
Y
```

```python
Y.value_counts()

```

# splitting the data into training and testing

<!-- grade_id: cell-50f6c46690e702a1, points: 0, grade: false, solution: true -->
```python
from sklearn.model_selection import train_test_split

#1 points
# do the train test split, tested with random_state= = 0
# you will need to make sure your data is a numpy array
X_train, X_test, Y_train, Y_test = 0,0,0,0
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-76833b6f347ba193, points: 1, grade: true, solution: false -->
```python

assert X_train.shape == (4457,)
assert X_test.shape == (1115,)
assert Y_train.shape == (4457,)
assert Y_test.shape == (1115,)
```

# Vectorization using Bag of Words

<!-- grade_id: cell-111f82315e4a5795, points: 0, grade: false, solution: true -->
```python
from sklearn.feature_extraction.text import CountVectorizer  #tf_idf
# 3 points
# use CountVectorizer to vectorize the data

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-fde26c3654f844d5, points: 3, grade: true, solution: false -->
```python
assert len(X_test[0]) == 29
assert len(X_train[0]) == 36
assert len(X_train[3]) == 21
assert len(X_test[3]) == 58
```

```python
X_train_vect
```

```python
X_test_vect
```

# Lets train our Model - Naive Bayes

<!-- grade_id: cell-f348d6e519765361, points: 0, grade: false, solution: true -->
```python
from sklearn.naive_bayes import GaussianNB
# 3 points

#train your model

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-cfaf9972f90a81a1, points: 3, grade: true, solution: false -->
```python
model.classes_
assert model.classes_.shape == (2,)
assert model.classes_[0] == 0
assert model.classes_[1] == 1
```

<!-- grade_id: cell-79a39702b247d13d, points: 0, grade: false, solution: true -->
```python
from sklearn.metrics import accuracy_score, confusion_matrix

#7 points

# predict and get the accuracy on test as well as the confusion matrix

acc = 0

conf_matrix = 0
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-3349b6b934a43864, points: 7, grade: true, solution: false -->
```python
assert Y_pred.shape == (1115,)
assert acc >= 80
assert conf_matrix[0][0] >= 3*conf_matrix[0][1]
assert conf_matrix[1][1] >= 3*conf_matrix[1][0]
```