# Exercise 10: Hierarchical clustering of the grain data

In the video, you learnt that the SciPy `linkage()` function performs hierarchical clustering on an array of samples.  Use the `linkage()` function to obtain a hierarchical clustering of the grain samples, and use `dendrogram()` to visualize the result.  A sample of the grain measurements is provided in the array `samples`, while the variety of each grain sample is given by the list `varieties`.


**Step 1:** Import:

 + `linkage` and `dendrogram` from `scipy.cluster.hierarchy`.
 + `matplotlib.pyplot` as `plt`.
 + `pandas`
 + `load_iris` and `train_test_split`

```python
import pandas as pd
import numpy as np

from sklearn.datasets import load_iris
from sklearn.model_selection import train_test_split
from scipy.cluster.hierarchy import linkage, dendrogram

import matplotlib.pyplot as plt

```

**Step 2:** Load the dataset _(done for you)_.

```python

iris_data = load_iris()

```

<!-- grade_id: cell-5b5863d4b8ed6db7, points: 0, grade: false, solution: true -->
```python
#Create pd DF the same way you did the first day
#Create the class and target columns and remove de cm from the columns names

def create_df():
    
    # YOUR CODE HERE
    raise NotImplementedError()
df_iris = create_df()
df_iris.sample(n=10)
```

<!-- grade_id: cell-632fd0e11b776633, points: 2, grade: true, solution: false -->
```python
assert df_iris['sepal length'].shape == (150,)
assert df_iris['sepal width'].shape == (150,)
assert df_iris['petal length'].shape == (150,)
assert df_iris['petal width'].shape == (150,)
assert df_iris['target'].shape == (150,)
assert df_iris['class'].shape == (150,)
```

<!-- grade_id: cell-1b1da2fc006dd2f6, points: 0, grade: false, solution: true -->
```python

#From our data we will only get variable columns as np.array
x = 0
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-81d6796d99e92aa3, points: 0, grade: true, solution: false -->
```python
assert x.shape == (150,4)
```

**Step 3:** Perform hierarchical clustering on `samples` using the `linkage()` function with the `method='complete'` keyword argument. Assign the result to `mergings`.

<!-- grade_id: cell-2a4681a4e740a6fe, points: 0, grade: false, solution: true -->
```python
mergings = 0 
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-bd07f2be641ae1c5, points: 3, grade: true, solution: false -->
```python
assert mergings.shape == (149,4)
```

**Step 4:** Plot a dendrogram using the `dendrogram()` function on `mergings`, specifying the keyword arguments `labels=varieties`, `leaf_rotation=90`, and `leaf_font_size=6`.  Remember to call `plt.show()` afterwards, to display your plot.

<!-- grade_id: cell-8ba6a5e72836aa92, points: 0, grade: false, solution: true -->
```python
dn = 0
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-bc0248a8151d10ab, points: 3, grade: true, solution: false -->
```python
assert type(dn) == type(dendrogram(np.random.rand(15, 4)))
```

## K-MEANS

for this next part we will use the K-Means algorithm in order to cluster your data

**Step 1:** 
+ import `KMeans` from `sklearn.cluster`
+ loading our new datasets

```python
from sklearn.cluster import KMeans

df = pd.read_csv('./datasets/ch1ex1.csv')
points = df.values

new_df = pd.read_csv('./datasets/ch1ex2.csv')
new_points = new_df.values

```

**Step 2:** Using `KMeans()`, create a `KMeans` instance called `model` to find `3` clusters. To specify the number of clusters, use the `n_clusters` keyword argument


<!-- grade_id: cell-d4136dfc2ba753ba, points: 0, grade: false, solution: true -->
```python
model = 0
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-a6a59ad2afd63249, points: 1, grade: true, solution: false -->
```python
assert type(model)== type(KMeans())
```

**Step 4:** Use the `.fit()` method of `model` to fit the model to the array of points `points`.

<!-- grade_id: cell-9a699df5bd74992a, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

**Step 5:** Use the `.predict()` method of `model` to predict the cluster labels of `points`, assigning the result to `labels`.

<!-- grade_id: cell-fe0636d9d3113913, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-3e790b82f9bf37fc, points: 2, grade: true, solution: false -->
```python
assert labels[labels.argmax()] == 2
assert labels.shape == (300,)
```

<!-- grade_id: cell-917ce87ba793a518, points: 0, grade: false, solution: true -->
```python
# Make a function that returns 3 numpy arrays each one with the points associated for each class
#If the label is 0 they go into data_0
#If the label is 1 they go into data_1
#If the label is 2 they go into data_2
def separate_labels(labels, points):
    data_0 = []
    data_1 = []
    data_2 = []

    # YOUR CODE HERE
    raise NotImplementedError()
    return data_0,data_1,data_2


data_0,data_1,data_2 =separate_labels(labels, points)

print(data_0.shape)
print(data_1.shape)
print(data_2.shape)
```

<!-- grade_id: cell-9d19700f8e07d490, points: 2, grade: true, solution: false -->
```python
assert abs(data_0.shape[0] - 94) <= 20
assert abs(data_1.shape[0] - 95) <= 20
assert abs(data_2.shape[0] - 111) <= 20
```

<!-- grade_id: cell-6ae9d7e6dfb48eda, points: 0, grade: false, solution: true -->
```python
# plotting the data 

# YOUR CODE HERE
raise NotImplementedError()
```

**Step 7:** Use the `.predict()` method of `model` to predict the cluster labels of `new_points`, assigning the result to `new_labels`.  Notice that KMeans can assign previously unseen points to the clusters it has already found!

<!-- grade_id: cell-e1c6be8c2b6c48ab, points: 0, grade: false, solution: true -->
```python

# YOUR CODE HERE
raise NotImplementedError()

```

<!-- grade_id: cell-33a135652cf843e7, points: 2, grade: true, solution: false -->
```python
assert new_labels[new_labels.argmax()] == 2
assert new_labels.shape == (100,)
```

<!-- grade_id: cell-91d525783ae0a408, points: 0, grade: false, solution: true -->
```python
#separate the data by labels like we did before

new_0 = []
new_1 = []
new_2 = []

# YOUR CODE HERE
raise NotImplementedError()

```

<!-- grade_id: cell-bbb9a9b41cb01987, points: 2, grade: true, solution: false -->
```python
assert abs(new_0.shape[0] - 34) <= 5
assert abs(new_1.shape[0] - 34) <= 5
assert abs(new_2.shape[0] - 34) <= 5
```

<!-- grade_id: cell-f0f51c7f91e8b4df, points: 0, grade: false, solution: true -->
```python
#plot the new labels


# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-fb66bc887cd85537, points: 0, grade: false, solution: true -->
```python
#plot all the points together

# YOUR CODE HERE
raise NotImplementedError()
```