# Linear Regression From Scratch
In this notebook you will implement linear regression using the least squares method.

The notebook will guide you in the general steps. You may use only numpy, pandas, and matplotlib for this exercise.

#### Math Background
The task is to solve linear regression for the data set of ```x``` and ```y```. That is, fit a line y to the data such that ```y = a + bx + e```. Where a and b are coefficents and e is an error term.

We know that ```b = SUM ( xi - x_mean) * (yi - y_mean) / SUM (xi - x_mean)^2``` where ```xi```, and ```yi``` are the indivdual observations and ```x_mean```, ```y_mean``` are means of all ```xi``` and ```yi```.

So a reasonable solution then for a fit is ```a = y_mean - b * x_mean```.


#### Implementation Steps
1. Load the file ```reg_data.csv```
2. Display the data with a scatter plot. Write a markdown cell and comment on the structure of the data and what kind of line you think will fit the data. 
3. Implement a function of the least squares method.
4. Plot the predictions from your function.
5. Comment on the output of your function. Does it make sense? Can it be made better? If yes, how?

```python
#import the minimum packages
import numpy as np
import pandas as pd
import matplotlib.pyplot as plt
from sklearn.model_selection import train_test_split
```

<!-- grade_id: cell-4104816f50e4844b, points: 0, grade: false, solution: true -->
```python
#load and inspect the data
df = 0
# YOUR CODE HERE
raise NotImplementedError()

```

<!-- grade_id: cell-3676b87d8fc79bbd, points: 1, grade: true, solution: false -->
```python
assert df.shape == (100,2)
```

<!-- grade_id: cell-7f576e50392bc5e9, points: 0, grade: false, solution: true -->
```python
#plot the data (scatter)
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-54891567ed3179cb, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-161d857ab77810d3, points: 1, grade: true, solution: false -->
```python
assert x_train.shape == (80,)
assert x_test.shape == (20,)
assert y_train.shape == (80,)
assert y_test.shape == (20,)
```

<!-- grade_id: cell-d5e750d407949b20, points: 0, grade: false, solution: true -->
```python
#implement a least squares function to find a, b

x_mean = 0
y_mean = 0
b_upper = 0
b_lower = 0
b = 0
a = 0

# YOUR CODE HERE
raise NotImplementedError()
print(b)
print(a)
```

<!-- grade_id: cell-c7fea6e45dc92701, points: 5, grade: true, solution: false -->
```python
assert abs(a - 7.72) <=0.03
assert abs(b - 1.32) <=0.03
```

<!-- grade_id: cell-721384e1a2152077, points: 0, grade: false, solution: true -->
```python
line = 0
x = np.array([])
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-71f1201b79ff7e44, points: 5, grade: true, solution: false -->
```python
assert abs(a +3*b - 11.69) <= 0.05
assert abs(a +100*b - 140) <= 1
```

<!-- grade_id: cell-e8d9edf92254aa00, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-8b80d3f35ea3b407, points: 0, grade: false, solution: true -->
```python
#Classify your test data in to classes
#if the Y value for a certain X is lower than the line then the class is 0
class_0 = []
class_1 = []

# YOUR CODE HERE
raise NotImplementedError()
        
class_0 = np.array(class_0)
class_1 = np.array(class_1)
print(class_0.shape)
print(class_1.shape)
```

<!-- grade_id: cell-73e31e3494663d82, points: 5, grade: true, solution: false -->
```python
assert 9 < class_0.shape[0] < 13
assert 7 < class_1.shape[0] < 11
```

<!-- grade_id: cell-db81fd9fa3a21bdf, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```